Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# I. $LU$ - разложение квадратной матрицы



Рассмотрим наивную реализацию LU - разложения.

Заметим, что мы используем массивы numpy для представления матриц. [Не используйте 'np.matrix'].


```python
import numpy as np

def diy_lu(a):
    """Создает LU - разложение матрицы `a`.
    
    Наивное LU - разложение: работает столбец за столбцом, накапливает элементарные треугольные матрицы.
    Без выбора главного элемента.
    """
    N = a.shape[0]
    
    u = a.copy()
    L = np.eye(N)
    for j in range(N-1):
        lam = np.eye(N)
        gamma = u[j+1:, j] / u[j, j]
        lam[j+1:, j] = -gamma
        u = lam @ u

        lam[j+1:, j] = gamma
        L = L @ lam
    return L, u
```

```python
# Теперь сгенерируем матрицу полного ранга и протестируем наивное разложение.
import numpy as np

N = 6
a = np.zeros((N, N), dtype=float)
for i in range(N):
    for j in range(N):
        a[i, j] = 3. / (0.6*i*j + 1)

L,U = diy_lu(a)

np.linalg.matrix_rank(a)
```

```python
# Настройка вывода чисел с плавающей точкой для большей ясности
np.set_printoptions(precision=3)
```

```python
L, u = diy_lu(a)
print(L, "\n")
print(u, "\n")

# Быстрый тест на адекватность: L @ U должна быть равна изначальной матрице с точностью до ошибок округления.
print(a-L@u)
```

# II. Необходимость выбора главного элемента

Давайте немного подправим матрицу, изменив в ней один элемент:

```python
a1 = a.copy()
a1[1, 1] = 3
print(a1)
```

Результирующая матрица имеет полный ранг, но наивное LU - разложение не работает.

```python
np.linalg.matrix_rank(a1)
```

```python
l, u = diy_lu(a1)

print(l, u, sep='\n')
```

### Тест II.1

Для того, чтобы наивное LU - разложение работало необходимо чтобы все лидирующие миноры матрицы были отличны от нуля. Проверьте, выполнено ли это требование для двух матриц a и a1.

(20% оценки)



<!-- grade_id: cell-cd7f666a9fa51d35, points: 0, grade: false, solution: true -->
```python
def minor(a):
    ''' Параметры
    ----------
    a : Квадратная матрица с действительными коэффициентами   
    Возвращает
    -------
    answer : boolean, True если все лидирующие миноры ненулевые.'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-212d71a3e76aea58, points: 2, grade: true, solution: false -->
```python
assert minor(a)==True
assert minor(a1)==False
```

### Тест II.2

Модифицируйте алгоритм diy_lu, чтобы осуществлять выбор главного элемента в столбцах. Для контроля выбора можете использовать матрицу перестановок или массив замен.

(40% оценки)

Напишите функию, воссоздающую изначальную матрицу из разложения. Протестируйте свой алгоритм на матрицах a и a1.

(40% оценки)

<!-- grade_id: cell-703a27f2e76339e7, points: 0, grade: false, solution: true -->
```python
def diy_lu_mod(a):
    ''' Parameters
    ----------
    a : Square matrix with float
       Coefficients
       
    Returns
    -------
    P, L, U : where P is a permutation matrix, 
    L lower triangular with unit diagonal elements, and U upper triangular.'''
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-eacc7aa962410524, points: 4, grade: true, solution: false -->
```python
P, L, U = diy_lu_mod(a)
P1, L1, U1 = diy_lu_mod(a1)
assert np.allclose([a,a1], [P@L@U, P1@L1@U1]), "Decomposition doesn't reconstruct the initial matrix"
```

<!-- grade_id: cell-b4a0e1ab6ef0038b, points: 4, grade: true, solution: false -->
```python
from scipy.linalg import lu
lu_true = [lu(a), lu(a1)]
assert np.allclose(lu_true, [[P,L,U],[P1,L1,U1]]), 'Decomposition is wrong'
```