Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Задачи по NumPy

```python
import numpy as np
import matplotlib.pyplot as plt
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-2-e0e1492b7973> in <module>
      1 import numpy as np
----> 2 import matplotlib.pyplot as plt

ModuleNotFoundError: No module named 'matplotlib'
```

1. Напишите функцию, возвращающую нулевой вектор длины 10, у которого пятая компонента единица.

<!-- grade_id: cell-f65d513623e98cd2, points: 0, grade: false, solution: true -->
```python
def null_vec():
    # YOUR CODE HERE
    a = np.empty(12)
    return a
```

<!-- grade_id: cell-1f7c19ae31380e34, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert np.allclose( null_vec() , np.array([0., 0., 0., 0., 1., 0., 0., 0., 0., 0.]))
assert str(type(null_vec()))=="<class 'numpy.ndarray'>"
###END HIDDEN TESTS
```

Output:
```
/opt/tljh/user/lib/python3.7/site-packages/numpy/core/numeric.py in isclose(a, b, rtol, atol, equal_nan)
   2288     yfin = isfinite(y)
   2289     if all(xfin) and all(yfin):
-> 2290         return within_tol(x, y, atol, rtol)
   2291     else:
   2292         finite = xfin & yfin

/opt/tljh/user/lib/python3.7/site-packages/numpy/core/numeric.py in within_tol(x, y, atol, rtol)
   2274     def within_tol(x, y, atol, rtol):
   2275         with errstate(invalid='ignore'):
-> 2276             return less_equal(abs(x-y), atol + rtol * abs(y))
   2277 
   2278     x = asanyarray(a)

ValueError: operands could not be broadcast together with shapes (12,) (10,)
```

2.Переверните вектор: [1, 2, 3]->[3,2,1]

<!-- grade_id: cell-9b4f47aa0cb85467, points: 0, grade: false, solution: true -->
```python
def reverse_vec(v):
    # YOUR CODE HERE
    raise NotImplementedError()
    return v_rev
```

<!-- grade_id: cell-f57d3b785304cbc2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
np.random.seed()
vec = np.random.rand(10)
assert np.allclose(reverse_vec(vec) , vec[::-1])
###END HIDDEN TESTS
```

Output:
```
<ipython-input-6-46f1a1ba19f9> in <module>
      2 np.random.seed()
      3 vec = np.random.rand(10)
----> 4 assert np.allclose(reverse_vec(vec) , vec[::-1])
      5 ###END HIDDEN TESTS

<ipython-input-5-ae3c2e09abcc> in reverse_vec(v)
      1 def reverse_vec(v):
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()
      4     return v_rev

NotImplementedError:
```

3.Создайте матрицу 3x3, в которой величины будут идти по порядку от 0 до 8: 

\begin{equation*} 
\begin{pmatrix}
0 & 1 & 2 \\
3 & 4 & 5 \\
6 & 7 & 8 
\end{pmatrix}
\end{equation*}

<!-- grade_id: cell-17bec3567fc74cc1, points: 0, grade: false, solution: true -->
```python
def matrix():
    # YOUR CODE HERE
    raise NotImplementedError()
    return m
matrix()
```

Output:
```
<ipython-input-7-13477530ef5d> in <module>
      3     raise NotImplementedError()
      4     return m
----> 5 matrix()

<ipython-input-7-13477530ef5d> in matrix()
      1 def matrix():
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()
      4     return m
      5 matrix()

NotImplementedError:
```

<!-- grade_id: cell-3b80759ec55892f3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert np.allclose(matrix() , np.arange(9).reshape((3, 3)))
###END HIDDEN TESTS
```

Output:
```
<ipython-input-8-f68ecf5ffba3> in <module>
      1 ###BEGIN HIDDEN TESTS
----> 2 assert np.allclose(matrix() , np.arange(9).reshape((3, 3)))
      3 ###END HIDDEN TESTS

<ipython-input-7-13477530ef5d> in matrix()
      1 def matrix():
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()
      4     return m
      5 matrix()

NotImplementedError:
```

4. Постройте матрицу размером 8х8 по принципу шахматной доски.

<!-- grade_id: cell-d3656b24b46b803b, points: 0, grade: false, solution: true -->
```python
def checkerboard(w, b):
    """Создает матрицу по принципу шахматной доски.
    Parameters
    ----------
    w: Символ, обозначающий белую клетку
    b: Символ, обозначающий черную клетку
    
    Returns
    -------
    m: Матрица 8х8.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return m
```

<!-- grade_id: cell-92e793300667d965, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
black = np.random.randint(5)
white = np.random.randint(5)
assert np.allclose(checkerboard(white, black) , np.array([black if abs((i//8)%2-i%2)==1 else white for i in range(64)]).reshape(8,8))
###END HIDDEN TESTS
```

Output:
```
<ipython-input-10-4aa7c6c3d4ef> in <module>
      2 black = np.random.randint(5)
      3 white = np.random.randint(5)
----> 4 assert np.allclose(checkerboard(white, black) , np.array([black if abs((i//8)%2-i%2)==1 else white for i in range(64)]).reshape(8,8))
      5 ###END HIDDEN TESTS

<ipython-input-9-9da051f048e2> in checkerboard(w, b)
     11     """
     12     # YOUR CODE HERE
---> 13     raise NotImplementedError()
     14     return m

NotImplementedError:
```

5. Напишите функцию, которая переводит декартовы координаты в полярные.

<!-- grade_id: cell-b190ffcc2e3da3a7, points: 0, grade: false, solution: true -->
```python
def cartesian_to_polar(c):
    """Переводит массив, состоящий из 10 пар координат xy 
    в массив координат r phi.
    
    Parameters
    ----------
    c: Array of int.
    Массив размерами 10х2, каждая строка -- декартовы координаты точки. 
    
    Returns
    -------
    p: Array of int. 
    Аналогичный массив из пар полярных координат. В первом столбике r.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return p
```

<!-- grade_id: cell-a4bb448a7328b920, points: 2, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
cart = np.array([np.random.randint(-10,10) for i in range(20)]).reshape(10, 2)
polar = cartesian_to_polar(cart)
assert np.allclose(polar[:, 0]*np.sin(polar[:, 1]), cart[:,1])
assert np.allclose(polar[:, 0]*np.cos(polar[:, 1]), cart[:,0])
###END HIDDEN TESTS
```

Output:
```
<ipython-input-12-5b8c8eaaca72> in <module>
      1 ###BEGIN HIDDEN TESTS
      2 cart = np.array([np.random.randint(-10,10) for i in range(20)]).reshape(10, 2)
----> 3 polar = cartesian_to_polar(cart)
      4 assert np.allclose(polar[:, 0]*np.sin(polar[:, 1]), cart[:,1])
      5 assert np.allclose(polar[:, 0]*np.cos(polar[:, 1]), cart[:,0])

<ipython-input-11-d81728c77c50> in cartesian_to_polar(c)
     14     """
     15     # YOUR CODE HERE
---> 16     raise NotImplementedError()
     17     return p

NotImplementedError:
```

6. Улитка Пасскаля -- плоская кривая, задаваемая уравнением в полярных координатах: $𝑟=𝑎+cos𝜙.$ 

Напишите функцию, определяющую координаты точек такой кривой.
(Подсказка: np.linspace создает одномерный массив из указанного количества элементов, значения которых равномерно распределенны внутри заданного интервала.)

<!-- grade_id: cell-bea2b11f79b37b02, points: 0, grade: false, solution: true -->
```python
def make_curve(n, a):
    '''Функция возвращает массив nх2, содержащий n пар 
    декартовых координат точек кривой Паскаля'''
    # YOUR CODE HERE
    raise NotImplementedError()
    return xy
```

<!-- grade_id: cell-28ce80f4d588a6d9, points: 2, grade: true, solution: false -->
```python
N = 50
A = 1
m = make_curve(N, A)
###BEGIN HIDDEN TESTS
check = m[np.random.choice(np.arange(N), 10)]
check = cartesian_to_polar(check)
assert np.allclose(check[:,0]-A, np.cos(check[:, 1]))
###END HIDDEN TESTS
```

Output:
```
<ipython-input-14-b0417b572b2a> in <module>
      1 N = 50
      2 A = 1
----> 3 m = make_curve(N, A)
      4 ###BEGIN HIDDEN TESTS
      5 check = m[np.random.choice(np.arange(N), 10)]

<ipython-input-13-d7200ca46e99> in make_curve(n, a)
      3     декартовых координат точек кривой Паскаля'''
      4     # YOUR CODE HERE
----> 5     raise NotImplementedError()
      6     return xy

NotImplementedError:
```

Постройте ее на плоскости xy для разных значений a.

<!-- grade_id: cell-5589004ce9d6a2fa, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-15-15b94d1fa268> in <module>
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```