Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Write a function which solves a quadratic equation, $x^2 + b x + c = 0$. Your function must receive two values for coefficients and return two real or complex roots. Always return two roots, even if the equation has a single double root. 

<!-- grade_id: cell-2032efc64a327acc, points: 0, grade: false, solution: true -->
```python
from math import sqrt
from cmath import sqrt as c_sqrt

def solve_quad(b, c):
    ''' Parameters
    ----------
    b, c : float
       Coefficients
       
    Returns
    -------
    x1, x2 : float or complex
        Roots.
    '''

    if b==0 and c==0:
      return 0, 0

    d = b*b-4*c    
    
    if d==0:
      x1 = (-b)/2
      x2 = (-b)/2
      print(x1, x2)
      return x1, x2

    if d > 0:    
      x1=-(2*c)/(sqrt(b*b-4*c)+b)
      x2=(-b-sqrt(d))/2
      print(x1, x2)
      return x1, x2
    else:
      x1=-(2*c)/(c_sqrt(b*b-4*c)+b)
      x2=(-b-c_sqrt(d))/2
      print(x1, x2)
      return x1, x2
```

<!-- grade_id: cell-89cf7badb387b0e0, points: 5, grade: true, solution: false -->
```python
#
# Your implementation should pass tests in this cell.
#
# Do not remove or alter this cell. You may run it, but do not remove or alter it.
# Your changes to this cell will be ignored on grading.
# There will be additional tests.
#
from numpy.testing import assert_allclose

variants = [{'b': 4.0, 'c': 3.0},
            {'b': 2.0, 'c': 1.0},
            {'b': 0.5, 'c': 4.0},
            {'b': 1e10, 'c': 3.0},
           ]

for var in variants:
    x1, x2 = solve_quad(**var)
    assert_allclose(x1*x2, var['c'], rtol=1e-15)
```

Output:
```
-1.0 -3.0
-1.0 -1.0
(-0.24999999999999994+1.9843134832984428j) (-0.25-1.984313483298443j)
-3e-10 -10000000000.0
```

<!-- grade_id: cell-0970e4298cb77486, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from numpy.testing import assert_allclose

hidden_variants = [{'b': 1e10, 'c': 3.0},
                   {'b': -1e10, 'c': 4.0},]
for var in hidden_variants:
    x1, x2 = solve_quad(**var)
    assert_allclose(x1*x2, var['c'], rtol=1e-15)
### END HIDDEN TESTS
```

Output:
```
-3e-10 -10000000000.0
<ipython-input-4-77c500eb0fc9> in <module>
      5                    {'b': -1e10, 'c': 4.0},]
      6 for var in hidden_variants:
----> 7     x1, x2 = solve_quad(**var)
      8     assert_allclose(x1*x2, var['c'], rtol=1e-15)
      9 ### END HIDDEN TESTS

<ipython-input-2-c4f19c84203f> in solve_quad(b, c)
     26 
     27     if d > 0:
---> 28       x1=-(2*c)/(sqrt(b*b-4*c)+b)
     29       x2=(-b-sqrt(d))/2
     30       print(x1, x2)

ZeroDivisionError: float division by zero
```