```python
import matplotlib.pyplot as plt
import numpy as np
```

# Часть I. Метод Ньютона нахождения корней уравнения.

Напишите алгоритм, который выполняет итерации Ньютона для заданной функции $f(x)$ с известной производной $f'(x)$. Ваша функция должна находить корни $f(x)$ с заданной точностью $\epsilon$.

<!-- grade_id: cell-e47469fedb8cc24a, points: 0, grade: false, solution: true -->
```python
def newton_iteration(f, fder, x0, eps=1e-5, maxiter=100):
    """Нахождение корней $f(x) = 0$ через итерации Ньютона.
    
    Parameters
    ----------
    f : callable
        Функция, корни которой мы хотим найти.
    fder : callable
        Производная `f`.
    x0 : float
        Начальное приближение итераций Ньютона.
    eps : float
        Заданная точность. 
        Алгоритм прекращает работу когда расстояние между последовательными приближениями меньше `eps`.
        По умолчанию 1e-5.
    maxiter : int
        Максимальное число итераций (по умолчанию 100).
        Алгоритм прекращается, когда число итераций достигает `maxiter`.
        Этот параметр нужен лишь для предотвращения бесконечного зацикливания.
    
    Returns
    -------
    x : float
        Найденное приближение к корню.
    niter : int
        Количество итераций.
    """
    ###BEGIN SOLUTION
    it = 0
    x, xprev = x0, x0 + 10
    while abs(x - xprev) > eps and it < maxiter:
        xprev = x
        x = x - f(x)/fder(x)
        it+=1
    return x, it
    ###END SOLUTION
```

<!-- grade_id: cell-0c23dfbb8447f5aa, points: 5, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

xx, nit = newton_iteration(lambda x: x**2 - 1, lambda x: 2.*x, x0=4)
assert_allclose(xx, 1.0, atol=1e-5)
assert nit < 10

xx, nit = newton_iteration(lambda x: x**2 - 1, lambda x: 2.*x, x0=-4)
assert_allclose(xx, -1.0, atol=1e-5)
assert nit < 10


from math import log, exp

xx, nit = newton_iteration(lambda x: exp(x)-2, lambda x: exp(x), x0=4, eps=1e-10)
assert_allclose(xx, log(2), atol=1e-10)
```

Используя функцию `newton_iteration` для функции $f(x) = x^2 - 1$, постройте зависимость логарифма ошибки нахождения корня методом Ньютона от числа итераций. Какова сходимость метода (линейная, квадратичная, иная)?

### Your code here.

```python

```

# Часть II. Модифицированный метод Ньютона для кратных корней.

Рассмотрим функцию, имеющую кратный корень (например, $f(x) = (x^2 - 1)^2$). Реализуйте модифицированную итерацию Ньютона, 

$$
x_{n+1} = x_{n} - m \frac{f(x_n)}{f'(x_n)}
$$

<!-- grade_id: cell-2f641d195eae18f0, points: 0, grade: false, solution: true -->
```python
def mod_newton(f, fder, x0, m, eps=1e-5, maxiter=100):
    """Нахождение корней $f(x) = 0$ через итерации Ньютона.
    
    Parameters
    ----------
    f : callable
        Функция, корни которой мы хотим найти.
    fder : callable
        Производная `f`.
    x0 : float
        Начальное приближение итераций Ньютона.
    eps : float
        Заданная точность. 
        Алгоритм прекращает работу когда расстояние между последовательными приближениями меньше `eps`.
        По умолчанию 1e-5.
    maxiter : int
        Максимальное число итераций (по умолчанию 100).
        Алгоритм прекращается, когда число итераций достигает `maxiter`.
        Этот параметр нужен лишь для преcдотвращения бесконечного зацикливания.
    
    Returns
    -------
    x : float
        Найденное приближение к корню.
    niter : int
        Количество итераций.
    """
    ###BEGIN SOLUTION
    it = 0
    x, xprev = x0, x0 + 10
    while abs(x - xprev) > eps and it < maxiter:
        xprev = x
        x = x - m*f(x)/fder(x)
        it+=1
    ###END SOLUTION
    return x, it
```

<!-- grade_id: cell-d2c8d42721e39f15, points: 5, grade: true, solution: false -->
```python
for m in [1, 2, 3, 4, 5, 6]:
    xx, nit = mod_newton(lambda x: (x**2 - 1)**4,
                         lambda x: 4*(x**2 - 1)**3 * 2 * x,
                         x0=2, m=m, maxiter=10000, eps=1e-9)
    assert_allclose(xx, 1.0, atol=1e-8)
    
### BEGIN HIDDEN TESTS
num_iters = []

for m in [1, 2, 3, 4, 5, 6]:
    xx, nit = mod_newton(lambda x: (x**2 - 1)**4,
                         lambda x: 4*(x**2 - 1)**3 * 2 * x,
                         x0=2, m=m, maxiter=10000, eps=1e-9)
    num_iters.append(nit)
    
assert num_iters[3] < 10
assert all(_ > 10 for _ in num_iters[:3])
assert all(_ > 10 for _ in num_iters[4:])
### END HIDDEN TESTS
```

Используйте модифицированную итерацию Ньютона для функции $f(x) = (x^2 - 1)^2$ со значениями  $m= 1, 2, 3, 4, 5, 6$. Сколько итераций требуется для достижения сходимости с фиксированным значением $\epsilon$? Согласуются ли ваши наблюдения с предсказанием о квадратичной сходимости метода при значении $m$ равном кратности корня, и линейной сходимости для остальных значений $m$?

### Your code here.

```python

```

# Часть III. Метод Ньютона для системы нелинейных уравнений.

Напишите алгоритм, который выполняет итерации Ньютона для заданной системы нелинейных уравнений $\vec{f}(x) = 0$ с известной матрицей Якоби $J(x)$. Ваша функция должна находить корни вектор-функции $\vec{f}(x)$ с заданной точностью $\epsilon$.

<!-- grade_id: cell-4580b8353412b2d3, points: 0, grade: false, solution: true -->
```python
def newton_system(F, jcb, X0, eps=1e-5, maxiter=100):
    """Нахождение корней $f(x) = 0$ через итерации Ньютона.
    
    Parameters
    ----------
    F : callable
        Вектор-функция системы, которую мы хотим решить.
    jcb : callable
        Производная `f`.
    X0 : array-like of floats, shape (n,)
        Начальное приближение итераций Ньютона.
    eps : float
        Заданная точность. 
        Алгоритм прекращает работу когда расстояние между последовательными приближениями меньше `eps`.
        По умолчанию 1e-5.
    maxiter : int
        Максимальное число итераций (по умолчанию 100).
        Алгоритм прекращается, когда число итераций достигает `maxiter`.
        Этот параметр нужен лишь для предотвращения бесконечного зацикливания.
    
    Returns
    -------
    X : array-like of floats, shape (n,)
        Найденное приближение к корню.
    niter : int
        Количество итераций.
    """
    ###BEGIN SOLUTION
    X = X0.copy()
    for niter in range(maxiter):
        A = jcb(X)
        b = F(X)
        dX = np.linalg.solve(A, -b)
        X += dX
        if np.linalg.norm(dX) < eps:
            break
    return X, niter
    ###END SOLUTION
```

<!-- grade_id: cell-32a522217549fd0c, points: 5, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

def func(X):
    return np.array([X[0]**2 - 2*X[1]**4 + 1, X[0] - X[1]**3 + 1])

def jac(X):
    return np.array([[2.*X[0], -8.*X[1]**3],
                       [1., -3.*X[1]**2]])

xx, nit = newton_system(func, jac, np.array([3., 2.]))
assert_allclose(xx, np.array([4.02879487, 1.71325223]), atol=1e-5)
assert nit < 10

xx, nit = newton_system(func, jac, np.asarray([1., 2.]))
assert_allclose(xx, np.asarray([-0.35119943,  0.86570596]), atol=1e-5)
assert nit < 10


### BEGIN HIDDEN TESTS
def func(X):
    x,y,z=X
    return np.array([x**2+y**2+z**2-25, x*y+y*z+z*x-5,x+y-3])

def jac(X):
    x,y,z=X
    return np.array([[2*x, 2*y, 2*z], [y+z, x+z, x+y], [1, 1, 0]])

xx, nit = newton_system(func, jac, np.asarray([4., 0., 3.]))
assert_allclose(xx, np.array([0.5*(3+np.sqrt(12*np.sqrt(35)-47)),
                              0.5*(3-np.sqrt(12*np.sqrt(35)-47)), 
                              np.sqrt(35)-3]), atol=1e-5)
assert nit < 10

xx, nit = newton_system(func, jac, np.asarray([1., 4., 3.]))
assert_allclose(xx, np.array([0.5*(3-np.sqrt(12*np.sqrt(35)-47)),
                              0.5*(3+np.sqrt(12*np.sqrt(35)-47)),
                              np.sqrt(35)-3]), atol=1e-5)
assert nit < 10
### END HIDDEN TESTS
```