Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Козловский Александр"
COLLABORATORS = "Петропавловский Владимир"
```

---

# I. $LU$ - разложение квадратной матрицы



Рассмотрим наивную реализацию LU - разложения.

Заметим, что мы используем массивы numpy для представления матриц. [Не используйте 'np.matrix'].


```python
import numpy as np

def diy_lu(a):
    """Создает LU - разложение матрицы `a`.
    
    Наивное LU - разложение: работает столбец за столбцом, накапливает элементарные треугольные матрицы.
    Без выбора главного элемента.
    """
    N = a.shape[0]
    
    u = a.copy()
    L = np.eye(N)
    for j in range(N-1):
        lam = np.eye(N)
        gamma = u[j+1:, j] / u[j, j]
        lam[j+1:, j] = -gamma
        u = lam @ u

        lam[j+1:, j] = gamma
        L = L @ lam
    return L, u
```

```python
# Теперь сгенерируем матрицу полного ранга и протестируем наивное разложение.
import numpy as np

N = 6
a = np.zeros((N, N), dtype=float)
for i in range(N):
    for j in range(N):
        a[i, j] = 3. / (0.6*i*j + 1)

L,U = diy_lu(a)

np.linalg.matrix_rank(a)
```

Output:
```
6
```

```python
# Настройка вывода чисел с плавающей точкой для большей ясности
np.set_printoptions(precision=3)
```

```python
L, u = diy_lu(a)
print(L, "\n")
print(u, "\n")

# Быстрый тест на адекватность: L @ U должна быть равна изначальной матрице с точностью до ошибок округления.
print(a-L@u)
```

Output:
```
[[1.    0.    0.    0.    0.    0.   ]
 [1.    1.    0.    0.    0.    0.   ]
 [1.    1.455 1.    0.    0.    0.   ]
 [1.    1.714 1.742 1.    0.    0.   ]
 [1.    1.882 2.276 2.039 1.    0.   ]
 [1.    2.    2.671 2.944 2.354 1.   ]] 

[[ 3.000e+00  3.000e+00  3.000e+00  3.000e+00  3.000e+00  3.000e+00]
 [ 0.000e+00 -1.125e+00 -1.636e+00 -1.929e+00 -2.118e+00 -2.250e+00]
 [ 0.000e+00  0.000e+00  2.625e-01  4.574e-01  5.975e-01  7.013e-01]
 [ 0.000e+00  2.220e-16  0.000e+00 -2.197e-02 -4.480e-02 -6.469e-02]
 [ 0.000e+00 -4.528e-16  0.000e+00  6.939e-18  8.080e-04  1.902e-03]
 [ 0.000e+00  4.123e-16  0.000e+00 -1.634e-17  0.000e+00 -1.585e-05]] 

[[ 0.000e+00  0.000e+00  0.000e+00  0.000e+00  0.000e+00  0.000e+00]
 [ 0.000e+00  0.000e+00  0.000e+00  0.000e+00  0.000e+00  0.000e+00]
 [ 0.000e+00  0.000e+00  0.000e+00 -2.220e-16  1.110e-16  1.665e-16]
 [ 0.000e+00  0.000e+00 -2.220e-16  5.551e-17  1.665e-16  1.665e-16]
 [ 0.000e+00  0.000e+00  1.110e-16  1.665e-16 -1.665e-16 -5.551e-17]
 
```

# II. Необходимость выбора главного элемента

Давайте немного подправим матрицу, изменив в ней один элемент:

```python
a1 = a.copy()
a1[1, 1] = 3
print(a1)
```

Output:
```
[[3.    3.    3.    3.    3.    3.   ]
 [3.    3.    1.364 1.071 0.882 0.75 ]
 [3.    1.364 0.882 0.652 0.517 0.429]
 [3.    1.071 0.652 0.469 0.366 0.3  ]
 [3.    0.882 0.517 0.366 0.283 0.231]
 [3.    0.75  0.429 0.3   0.231 0.188]]
```

Результирующая матрица имеет полный ранг, но наивное LU-разложение не работает.

```python
np.linalg.matrix_rank(a1)
```

Output:
```
6
```

```python
l, u = diy_lu(a1)

print(l, u, sep='\n')
```

Output:
```
[[nan nan nan nan nan nan]
 [nan nan nan nan nan nan]
 [nan nan nan nan nan nan]
 [nan nan nan nan nan nan]
 [nan nan nan nan nan nan]
 [nan nan nan nan nan nan]]
[[nan nan nan nan nan nan]
 [nan nan nan nan nan nan]
 [nan nan nan nan nan nan]
 [nan nan nan nan nan nan]
 [nan nan nan nan nan nan]
 [nan nan nan nan nan nan]]
<ipython-input-4-8159eb3fe7f6>:15: RuntimeWarning: divide by zero encountered in true_divide
  gamma = u[j+1:, j] / u[j, j]
<ipython-input-4-8159eb3fe7f6>:17: RuntimeWarning: invalid value encountered in matmul
  u = lam @ u
<ipython-input-4-8159eb3fe7f6>:20: RuntimeWarning: invalid value encountered in matmul
  L = L @ lam
<ipython-input-4-8159eb3fe7f6>:15: RuntimeWarning: invalid value encountered in true_divide
  gamma = u[j+1:, j] / u[j, j]
```

### Тест II.1

Для того, чтобы наивное LU - разложение работало необходимо чтобы все лидирующие миноры матрицы были отличны от нуля. Проверьте, выполнено ли это требование для двух матриц a и a1.

<!-- grade_id: cell-cd7f666a9fa51d35, points: 0, grade: false, solution: true -->
```python
def minor(a):
    ''' Check if all leading minors are non-zero.
    
    Parameters
    ----------
    a : np.array
        2D array representing the square matrix
    
    Returns
    -------
    answer : bool
        True if all leading minors are non-zero
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-212d71a3e76aea58, points: 2, grade: true, solution: false -->
```python
assert minor(a)==True
assert minor(a1)==False
```

Output:
```
<ipython-input-12-3c97a2478811> in <module>
----> 1 assert minor(a)==True
      2 assert minor(a1)==False

<ipython-input-11-dc878fd6755a> in minor(a)
     13     '''
     14     # YOUR CODE HERE
---> 15     raise NotImplementedError()

NotImplementedError:
```

### Тест II.2

Модифицируйте алгоритм diy_lu, чтобы осуществлять выбор главного элемента в столбцах. Для контроля выбора можете использовать матрицу перестановок или массив замен.
Напишите функцию, воссоздающую изначальную матрицу из разложения. Протестируйте свой алгоритм на матрицах a и a1.

<!-- grade_id: cell-703a27f2e76339e7, points: 0, grade: false, solution: true -->
```python
def diy_lu_mod(A):
    '''Perform pivoted LU factorization of the input matrix. 
    
    Parameters
    ----------
    a : np.array
        2D array representing a square matrix with float entries
       
    Returns
    -------
    P, L, U : ndarrays
        factors. Here P is a permutation matrix, L is lower triangular
        with unit diagonal elements, and U upper triangular.
    '''
    U = A.copy()
    N = A.shape[0]
    L, P = np.eye(N), np.eye(N)
    for j in range(N-1):
        lam, r = np.eye(N), np.eye(N)
        row = np.argmax(abs(U[j:, j])) + j 
        if row != j:
            r[[j, row]] = r[[row, j]]
            P, U, L = r@P, r@U, r@L

        gamma = u[j+1:, j] / u[j, j]
        lam[j+1:, j] = -gamma
        U = lam @ U
        L = lam @ L
    
    L = P @ (np.linalg.inv(L))
    return P.T, L, U

```

<!-- grade_id: cell-eacc7aa962410524, points: 8, grade: true, solution: false -->
```python
P, L, U = diy_lu_mod(a)
P1, L1, U1 = diy_lu_mod(a1)

from numpy.testing import assert_allclose
assert_allclose(a, P @ L @ U, atol=1e-10)
assert_allclose(a1, P1 @ L1 @ U1, atol=1e-10 )

```

Output:
```
[... skipping hidden 2 frame]

c:\users\xiaomi\appdata\local\programs\python\python38-32\lib\site-packages\numpy\testing\_private\utils.py in func_assert_same_pos(x, y, func, hasval)
    738                                 % (hasval), verbose=verbose, header=header,
    739                                 names=('x', 'y'), precision=precision)
--> 740             raise AssertionError(msg)
    741         # If there is a scalar, then here we know the array has the same
    742         # flag as it everywhere, so we should return the scalar flag.

AssertionError: 
Not equal to tolerance rtol=1e-07, atol=1e-10

x and y nan location mismatch:
 x: array([[3.      , 3.      , 3.      , 3.      , 3.      , 3.      ],
       [3.      , 1.875   , 1.363636, 1.071429, 0.882353, 0.75    ],
       [3.      , 1.363636, 0.882353, 0.652174, 0.517241, 0.428571],...
 y: array([[nan, nan, nan, nan, nan, nan],
       [nan, nan, nan, nan, nan, nan],
       [nan, nan, nan, nan, nan, nan],...
```

```python

```