Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Анна Авилова"
COLLABORATORS = ""
```

---

# Интерполяция при помощи полиномов Лагранжа.

<!-- grade_id: cell-d459dee4d09527c4, points: 0, grade: false, solution: true -->
```python
import numpy as np

class LagrangeInterpolator:
    """Lagrange interpolating polynomial.
    
    Given a set of pairs ``(x_k, y_k)``, construct 
    a Lagrange polynomial ``f(x)``, such that
    
    .. math::

        f(x_k) = y_k   for k =0, ..., n-1
    
    Parameters
    ----------
    xk : array_like, shape(n,)
        Abscissas
    yk : array_like, shape(n,)
        Ordinates
    
    Attributes
    ----------
    __call__
    
    """
    def __init__(self, xk, yk):
        self.xk = np.asarray(xk, dtype=float)
        self.yk = np.asarray(yk, dtype=float)
        
    def __call__(self, x):
        """Evaluate the interpolator at a given point.
        
        Parameters
        ----------
        x : float
        
        Returns
        -------
        the value of the interpolator at ``x``.
        """
        summ = 0.0
        m = len(self.xk)
        for i in range(m):
            result = 1.0
            for j in range(m):
                if i==j:
                    continue
                num = x - self.xk[j]
                den = self.xk[i] - self.xk[j]
                result *= num / den
            summ += result * self.yk[i]
        return summ
```

<!-- grade_id: cell-02e34e05d317dfb2, points: 5, grade: true, solution: false -->
```python
def runge_func(x, a=25):
    return 1.0 / (1.0 + a*x**2)

xx = np.linspace(-2, 2, 21)
yy = runge_func(xx)

lagr = LagrangeInterpolator(xx, yy)

from numpy.testing import assert_allclose

assert_allclose(yy,
                [lagr(xval) for xval in xx],
                atol=1e-14)
```

Рассмотрим функцию Рунге $1/(1 + 25x^2)$. Интерполируйте данную функцию на интервале $x\in [-2, 2]$, используя полиномы Лагранжа степени $m$, где $m=3, 5, 7, 11$. Используйте равномерную сетку. Постройте результат интерполяции вместе с $f(x)$.

Повторите предыдущее задание используя узлы Чебышева. Сравните качество интерполяции на равномерной сетке и на узлах Чебышева.


```python
import matplotlib.pyplot as plt

t = np.linspace(-2, 2, 100)

fig, ax = plt.subplots(2, 2, figsize = (12, 8))
fig.suptitle('Equidistant nodes', fontsize = 16, y = 1)

for i,m in enumerate([3, 5, 7, 11]):
    
    ax[(i)//2, (i)%2].plot(t, runge_func(t), 'c', label = 'given function')
    nodes = np.linspace(-2, 2, m)
    lagr = LagrangeInterpolator(nodes, runge_func(nodes))
    
    ax[(i)//2, (i)%2].plot(nodes, runge_func(nodes), 'ro')
    ax[(i)//2, (i)%2].plot(t, [lagr(_) for _ in t], '--r', label = 'interpolation')
    ax[(i)//2, (i)%2].legend()
    ax[(i)//2, (i)%2].set_title(f'число узлов = {m}')
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-4-9e34136cefb8> in <module>
----> 1 import matplotlib.pyplot as plt
      2 
      3 t = np.linspace(-2, 2, 100)
      4 
      5 fig, ax = plt.subplots(2, 2, figsize = (12, 8))

ModuleNotFoundError: No module named 'matplotlib'
```

```python
def chebyshev_nodes(a, b, n):
    return [1/2*(a + b) + 1/2*(b - a)*np.cos(np.pi * (2*k - 1) / (2*n)) for k in range(1, n + 1)]

t = np.linspace(-2, 2, 100)

fig, ax = plt.subplots(2, 2, figsize = (12, 8))
fig.suptitle('Сhebyshev nodes', fontsize = 16, y = 1)

for i,m in enumerate([3, 5, 7, 11]):
    
    ax[(i)//2, (i)%2].plot(t, runge_func(t), 'r', label = 'given function')
    nodes = np.array(chebyshev_nodes(-2, 2, m))
    lagr = LagrangeInterpolator(nodes, runge_func(nodes))
    
    ax[(i)//2, (i)%2].plot(nodes, runge_func(nodes), 'ro')
    ax[(i)//2, (i)%2].plot(t, [lagr(_) for _ in t], '--c', label = 'interpolation')
    ax[(i)//2, (i)%2].legend()
    ax[(i)//2, (i)%2].set_title(f'число узлов: = {m}')
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-5-bea52ad85379> in <module>
      4 t = np.linspace(-2, 2, 100)
      5 
----> 6 fig, ax = plt.subplots(2, 2, figsize = (12, 8))
      7 fig.suptitle('Сhebyshev nodes', fontsize = 16, y = 1)
      8 

NameError: name 'plt' is not defined
```