Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Timofey Vilkov"
COLLABORATORS = ""
```

---

Write a function which solves a quadratic equation, $x^2 + b x + c = 0$. Your function must receive two values for coefficients and return two real or complex roots. Always return two roots, even if the equation has a single double root. 

<!-- grade_id: cell-2032efc64a327acc, points: 0, grade: false, solution: true -->
```python
from math import sqrt
from cmath import sqrt as c_sqrt

def solve_quad(b, c):
    if b > 5000*c:
        discriminant = b**2 - 4*c
        x1 = (-b - c_sqrt(discriminant)) / 2
        x2 = c/x1
        return x1, x2
    elif b != 0 and c != 0:
        discriminant = b**2 - 4*c
        x1 = (-b + c_sqrt(discriminant)) / 2
        x2 = (-b - c_sqrt(discriminant)) / 2
        return x1, x2
    elif b != 0 and c == 0:
        x1 = 0
        x2 = -b
        return x1, x2
    elif  b == 0 and c != 0:
        x1 = c_sqrt(c)
        x2 = -x1
        return x1, x2
    else:
        x1 = 0
        x2 = 0
    return 0, 0

```

<!-- grade_id: cell-89cf7badb387b0e0, points: 5, grade: true, solution: false -->
```python
#
# Your implementation should pass tests in this cell.
#
# Do not remove or alter this cell. You may run it, but do not remove or alter it.
# Your changes to this cell will be ignored on grading.
# There will be additional tests.
#
from numpy.testing import assert_allclose

variants = [{'b': 4.0, 'c': 3.0},
            {'b': 2.0, 'c': 1.0},
            {'b': 0.5, 'c': 4.0},
            {'b': 1e10, 'c': 3.0},
           ]

for var in variants:
    x1, x2 = solve_quad(**var)
    assert_allclose(x1*x2, var['c'], rtol=1e-15)
```

<!-- grade_id: cell-0970e4298cb77486, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from numpy.testing import assert_allclose

hidden_variants = [{'b': 1e10, 'c': 3.0},
                   {'b': -1e10, 'c': 4.0},]
for var in hidden_variants:
    x1, x2 = solve_quad(**var)
    assert_allclose(x1*x2, var['c'], rtol=1e-15)
### END HIDDEN TESTS
```

Output:
```
~/anaconda3/lib/python3.7/site-packages/numpy/testing/_private/utils.py in assert_allclose(actual, desired, rtol, atol, equal_nan, err_msg, verbose)
   1531     header = 'Not equal to tolerance rtol=%g, atol=%g' % (rtol, atol)
   1532     assert_array_compare(compare, actual, desired, err_msg=str(err_msg),
-> 1533                          verbose=verbose, header=header, equal_nan=equal_nan)
   1534 
   1535 

~/anaconda3/lib/python3.7/site-packages/numpy/testing/_private/utils.py in assert_array_compare(comparison, x, y, err_msg, verbose, header, precision, equal_nan, equal_inf)
    844                                 verbose=verbose, header=header,
    845                                 names=('x', 'y'), precision=precision)
--> 846             raise AssertionError(msg)
    847     except ValueError:
    848         import traceback

AssertionError: 
Not equal to tolerance rtol=1e-15, atol=0

Mismatched elements: 1 / 1 (100%)
Max absolute difference: 4.
Max relative difference: 1.
 x
```