## Task 1
Solve a quadratic equation, ð‘¥2+ð‘ð‘¥+ð‘=0 with discriminante. Write a formula of x1, x2.

<!-- grade_id: cell-8c916f5ac0225f9c, points: 1, grade: true, solution: true -->
Write your solution here

## Task 2
Solve a quadratic equation, $x^2 + b x + c = 0$.

<!-- grade_id: cell-2032efc64a327acc, points: 0, grade: false, solution: true -->
```python
def solve_quad(b, c):
    ''' Parameters
    ----------
    b, c : float
       Coefficients
       
    Returns
    -------
    x1, x2 : float or complex
    Roots.'''
    ### BEGIN SOLUTION
    if abs(b)>10**5:
        x1 = -c/b; x2=-b
        return x1,x2
    else:
        d = b ** 2 - 4 * c
    a = -b / 2 
    b = abs(d**0.5) / 2 
    ANS = [complex(a,-b), complex(a,b)] if d<0 else [float(a)+float(-b) , float(a)-float(-b)]
    return ANS[0], ANS[1]
    ### END SOLUTION
```

<!-- grade_id: cell-89cf7badb387b0e0, points: 1, grade: true, solution: false -->
```python
from numpy import allclose

variants = [{'b': 4.0, 'c': 3.0},
            {'b': 2.0, 'c': 1.0},
            {'b': 0.5, 'c': 4.0}]

for var in variants:
    x1, x2 = solve_quad(**var)
    assert x1*x2 == var['c']

```

<!-- grade_id: cell-0970e4298cb77486, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
hidden_variants = [{'b': 1e10, 'c': 3.0}, {'b': -1e10, 'c': 4.0},]
for var in hidden_variants:
    x1, x2 = solve_quad(**var)
    assert x1*x2 == var['c']
### END HIDDEN TESTS
```

```python

```