Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Morozov Alexandr"
COLLABORATORS = "Morozov Alexandr"
```

---

# Метод наименьших квадратов (линейный случай).

Рассмотрим функцию $y = f(x)$ определенной набором значений $y_0, y_1, \cdots, y_n$ в точках $x_0, x_1, \cdots, x_n$.

```python
import matplotlib.pyplot as plt
import numpy.linalg as la
import numpy.polynomial as pol
```

```python
import numpy as np

x = np.asarray([-1, -0.7, -0.43, -0.14, 0.14, 0.43, 0.71, 1, 1.29, 1.57, 1.86, 2.14, 2.43, 2.71, 3])
y = np.asarray([-2.25, -0.77, 0.21, 0.44, 0.64, 0.03, -0.22, -0.84, -1.2, -1.03, -0.37, 0.61, 2.67, 5.04, 8.90])

rng = np.random.default_rng(123)
x += 0.1 * rng.uniform(size=len(x))
y += 4 * rng.uniform(size=len(x))
```

### I. Найдите многочлен наилучшего приближения,

$$
P_m(x) = a_0 + a_1 x + \cdots + a_m x^m
$$

используя линейный метод наименьших квадратов.Для этого:

1. реализуйте алгоритм, создающий  "design matrix" `А`, используя $1, x, \cdots, x^m$ в качестве базисных функций.

2. постройте явно нормальную систему уравнений задачи наименьших квадратов при фиксированных значениях $m$.

3. решите систему уравнений, чтобы найти коэффициенты $P_m(x)$ для $m = 0, 1, 2, \dots$. Для решения задач линейной алгебры можете использовать либо библиотечные функции (`numpy.linalg.solve`), либо Вашу реализацию LU - разложение из 1 недели.

<!-- grade_id: cell-3c5b770ce3dc9362, points: 0, grade: false, solution: true -->
```python
import numpy as np

def design_matrix(x, m):
    """Construct the design matrix with monomials x**k for k=0..m-1"""
    M = np.empty((len(x), m), dtype=float)
    for i in range(m):
        M[:, i] = x**i
    return M
    
    raise NotImplementedError()
```

<!-- grade_id: cell-4eb9827586286964, points: 1, grade: true, solution: false -->
```python
assert design_matrix(x, 3).shape == (len(x), 3)
assert design_matrix(x, 11).shape == (len(x), 11)
```

<!-- grade_id: cell-ec39ee7d9ada0eb5, points: 0, grade: false, solution: true -->
```python
def lsq_poly(x, y, m):
    """Construct the LSQ polynomial of degree `m-1`.
    
    Parameters
    ----------
    x : array_like
        Sample points
    y : array_like
        Measured values
    m : int
        The number of coefficients of the LSQ polynomial
        (i.e. the degree of the polynomial is `m-1`)
        
    Returns
    -------
    p : callable 
        np.polynomial.Polynomial instance, representing the LSQ polynomial
        
    Examples
    --------
    >>> p = lsq_poly([1, 2, 3], [4, 5, 6], m=2)
    >>> p(np.array([1.5, 2.5]))
    array([4.5, 5.5])
    
    """
    x = np.asarray(x, dtype=float)
    y = np.asarray(y, dtype=float)
    if x.shape[0] != y.shape[0]:
        raise ValueError("Expect paired data.")
    if x.shape[0] <= m:
        raise ValueError("Cannot fit a degree-%s polynomial through %s points" % (m, x.shape[0]))
    
    M = design_matrix(x, m)
    
    c = la.inv(M.T @ M) @ M.T
    P = pol.Polynomial(c.dot(y))
    return P

    raise NotImplementedError()
```

<!-- grade_id: cell-c690d81357c56b85, points: 3, grade: true, solution: false -->
```python
import pytest
with pytest.raises(ValueError):
    lsq_poly([1, 2, 3], [1, 2], m=1)

with pytest.raises(ValueError):
    lsq_poly([1, 2, 3], [1, 2, 3], m=3)
    
# test that lsq_poly returns a callable object
# which support vectorized evaluation
p = lsq_poly(x, y, m=3)
assert len(p(x)) == len(x)
```

Вычислите число обусловленности матрицы `A` и матрицы левой части системы нормальных уравнений в зависимости $m$, и напечатайте их на экран.
Прокомментируйте полученные результаты.

<!-- grade_id: cell-ede76d4a7cff87d0, points: 0, grade: false, solution: true -->
```python
for m in range(1, 10):
    M = design_matrix(x, m)
    c_f = la.cond(M)
    c_s = la.cond(M.T @ M)
    print(c_f, c_s, "\n")
    
raise NotImplementedError()
```

Output:
```
1.0 1.0 

2.5436796556599033 6.470306190618083 

7.974123103608077 63.586639271496125 

33.65632525997138 1132.7482300049676 

149.21938644682731 22266.425291569114 

709.163740691876 502913.2111124388 

3627.8477937656953 13161279.615912767 

18346.198895485024 336583013.98491627 

102424.59756750132 10490798237.031221 

NotImplementedError                       Traceback (most recent call last)
<ipython-input-17-04a248b2a7a3> in <module>
      5     print(c_f, c_s, "\n")
      6 
----> 7 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-a57fd1f730a5ca62, points: 1, grade: true, solution: false -->
```python
# This cell is a test: keep it intact

```

### II. Многочлен наилучшего приближения.

Ясно, что бесконтрольно увеличивать степень аппроксимирующего многочлена довольно бессмысленно. Зададимся целью построить "оптимальный" многочлен как можно более низкой степени $m$.

Для нахождения оптимального значения $m$ используйте следующий критерий: возьмем $m = 1, 2, \dots$, и для каждого значения $m$ вычислим

$$
\sigma_m^2 = \frac{1}{n - m} \sum_{k=0}^n \left( P_m(x_k) - y_k \right)^2
$$

и возьмем такое значение $m$, при котором $\sigma_m$ стабилизируется или начинает расти.

<!-- grade_id: cell-5d7424589f59ac28, points: 0, grade: false, solution: true -->
```python
def sigma(x, y, m):
    r"""Compute $\sigma_m$."""
    
    l = len(x)
    p = lsq_poly(x, y, m)
    
    s = 1/(l - m)
    s = s * np.sum((p(x) - y) ** 2)
    return s

    raise NotImplementedError()
```

```python
for m in range(1, 15):
    print(m, sigma(x, y, m))
```

Output:
```
1 11.457382047310073
2 8.520081963394484
3 4.005564058238968
4 1.1356406880121819
5 1.1272670188279166
6 1.1208555131773144
7 1.0050279843985317
8 1.1003064743050361
9 1.2659515051676018
10 0.7701191223054554
11 0.1596500955253315
12 0.17430729774185805
13 0.14408059529043296
14 485.1797943075314
```

Постройте ваши полиномы $P_m(x)$ вместе с данными. Визуально сравните полиномы различных степеней. Согласуется ли визуальное сравнение с ожидаемым оптимальным значением $m$?

```python
import matplotlib.pyplot as plt
fig, ax = plt.subplots()
for m in range(1, 7):
    p = lsq_poly(x, y, m)
    
    xx = np.linspace(-1, 3.2, 51)
    ax.plot(xx, p(xx), '-', label='m = %s' % m)

ax.plot(x, y, 'o', ms=8, color='k')
plt.legend(loc='best')
plt.grid()
```

### III. МНК с использованием QR-разложения.

Для оптимального значения $m$ из предыдущей части решите задачу МНК, QR-разложение design матрицы. Для решения задач линейной алгебры можете использовать библиотечные функции (обратите внимание на `scipy.linalg.solve_triangular` и `scipy.linalg.qr` и т.д.) или код предыдущих недель.
Сравните результат с ответом, полученным при решении системы уравнений.

<!-- grade_id: cell-2f72bccc7b48651d, points: 0, grade: false, solution: true -->
```python
def lsq_qr(x, y, m):
    """Solve the LSQ problem via the QR decomp of the design matrix.

    Parameters
    ----------
    x : array_like
        Sample points
    y : array_like
        Measured values
    m : int
        The degree of the LSQ polynomial
        
    Returns
    -------
    p : callable 
        np.polynomial.Polynomial instance, representing the LSQ polynomial
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e84a63fbccd75fdf, points: 2, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

for m in range(1, 10):
    assert_allclose(lsq_qr(x, y, m=m).coef,
                    lsq_poly(x, y, m=m).coef,
                    atol=1e-14)
```

<!-- grade_id: cell-0dec472419aee2c0, points: 3, grade: true, solution: false -->
```python
# This cell spot-tests several polynomials to make sure that they are worse than your ``best'' polynomial.


```

```python

```