Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

```python
import numpy as np

from numpy.testing import assert_allclose
```


### Часть I. Постройте отражение Хаусхолдера для вектора.

Дан вектор $\mathbf{x}$ и плоскость, заданная вектором нормали $\mathbf{u}$. Преобразование Хаусхолдера отражает $\mathbf{x}$ относительно плоскости.

Матрица преобразований Хаусхолдера:
$$ \mathbf{H} = \mathbf{1} - 2 \mathbf{u} \mathbf{u}^T $$

Если даны два вектора $\mathbf{x}$ и $\mathbf{y}$ одинаковой длины, поворот, преобразующий $\mathbf{x}$ в $\mathbf{y}$ называется преобразованием Хаусхолдера с
$$ \mathbf{u} = \frac{\mathbf{x} - \mathbf{y}}{\left|\mathbf{x} - \mathbf{y}\right|} $$

Напишите функцию, преобразующую заданный вектор $\mathbf{x} = (x_1, \dots, x_n)$ в $\mathbf{y} = (\left|\mathbf{x}\right|, 0, \dots, 0)^T$, используя преобразование Хаусхолдера.


<!-- grade_id: cell-d8d82542e794c193, points: 0, grade: false, solution: true -->
```python
def householder(vec):
    """Создает отражение Хаусхолдера, преобразующее 2-ю, 3-ю и т.д. компоненты вектора vec в нули.
    
    Parameters
    ----------
    vec : array-like of floats, shape (n,)
        Введенный вектор
    
    Returns
    -------
    outvec : array of floats, shape (n,)
        Преобразованный вектор, причем ``outvec[1:]==0`` и ``|outvec| == |vec|``
    H : array of floats, shape (n, n)
        Ортогональная матрица отражений Хаусхолдера
    """



    vec = np.asarray(vec, dtype=float)
    if vec.ndim != 1:
        raise ValueError("vec.ndim = %s, expected 1" % vec.ndim)

    if np.size(vec)==1 and np.linalg.norm(vec)<1e-8:
      return([0],[[1]])

    x = vec
    y = np.zeros(np.size(vec))
    y[0] = np.linalg.norm(vec)
    u = (x - y)/np.linalg.norm(x - y)
    H=np.eye(np.size(vec))-2*(u[np.newaxis, :].T @ u[np.newaxis, :] )

    return(y, H)
```

```python
    vec = np.array([1, 2, 3])

    vec = np.asarray(vec, dtype=float)
    if vec.ndim != 1:
        raise ValueError("vec.ndim = %s, expected 1" % vec.ndim)

    x = vec
    y = np.zeros(np.size(vec))
    y[0] = np.linalg.norm(vec)
    #y = y[np.newaxis, :].T
    print(x - y)
    print(np.linalg.norm(x - y))
    u = (x - y)/np.linalg.norm(x - y)
    print(u)
    print(u[np.newaxis, :].T)
    
    print(  u[np.newaxis, :].T @ u[np.newaxis, :] )
    print(  2*(u[np.newaxis, :].T @ u[np.newaxis, :] ))
    print(1-  2*(u[np.newaxis, :].T @ u[np.newaxis, :] ))

    H=np.eye(3)-2*(u[np.newaxis, :].T @ u[np.newaxis, :] )
    print(H)

    print(H @ x)
```

Output:
```
[-2.74165739  2.          3.        ]
4.529534769317056
[-0.60528454  0.44154645  0.66231968]
[[-0.60528454]
 [ 0.44154645]
 [ 0.66231968]]
[[ 0.36636938 -0.26726124 -0.40089186]
 [-0.26726124  0.19496327  0.2924449 ]
 [-0.40089186  0.2924449   0.43866735]]
[[ 0.73273876 -0.53452248 -0.80178373]
 [-0.53452248  0.38992654  0.5848898 ]
 [-0.80178373  0.5848898   0.87733471]]
[[0.26726124 1.53452248 1.80178373]
 [1.53452248 0.61007346 0.4151102 ]
 [1.80178373 0.4151102  0.12266529]]
[[ 0.26726124  0.53452248  0.80178373]
 [ 0.53452248  0.61007346 -0.5848898 ]
 [ 0.80178373 -0.5848898   0.12266529]]
[ 3.74165739e+00  0.00000000e+00 -1.11022302e-16]
```

Протестируйте свою функцию на следующих примерах:

<!-- grade_id: cell-7364c2b9de6bb44c, points: 3, grade: true, solution: false -->
```python
v = np.array([1, 2, 3])
v1, h = householder(v)

assert_allclose(h @ v1, v, atol=1e-14)
assert_allclose(h @ v, v1, atol=1e-14)

assert_allclose(v1[1:], 0, atol=1e-14)

assert_allclose(h @ h.T, np.eye(3), atol=1e-14)

```

### Part II. Вычисление $\mathrm{QR}$ - разложения матрицы.

Дана прямоугольная $m\times n$ матрица $\mathbf{A}$. Выполните отражение Хаусхолдера матрицы $\mathbf{H}_1$, преобразующее первый столбец матрицы $\mathbf{A}$ (назовем результат $\mathbf{A}^{(1)}$)

$$ 
\mathbf{H}_1 \mathbf{A} %\begin{pmatrix} \times &amp; \times &amp; \times &amp; \dots &amp; \times \\ 0 &amp; \times &amp; \times &amp; \dots &amp; \times \\ 0 &amp; \times &amp; \times &amp; \dots &amp; \times \\ &amp;&amp; \dots&amp;&amp; \\ 0 &amp; \times &amp; \times &amp; \dots &amp; \times \\ \end{pmatrix}
%
\equiv \mathbf{A}^{(1)}\;. 
$$

Теперь рассмотрим нижнюю правую часть матрицы $\mathbf{A}^{(1)}$ и выполним преобразование Хаусхолдера, действующее на 2 столбец $\mathbf{A}$:
$$
\mathbf{H}_2 \mathbf{A}^{(1)} % \begin{pmatrix} \times &amp; \times &amp; \times &amp; \dots &amp; \times \\ 0 &amp; \times &amp; \times &amp; \dots &amp; \times \\ 0 &amp; 0 &amp; \times &amp; \dots &amp; \times \\ &amp;&amp; \dots&amp;&amp; \\ 0 &amp; 0 &amp; \times &amp; \dots &amp; \times \\ \end{pmatrix}
%
\equiv \mathbf{A}^{(2)} \;. $$

Повторяя процесс $n$ раз, получим
$$
\mathbf{H}_{n} \cdots \mathbf{H}_2 \mathbf{H}_1 \mathbf{A} = \mathbf{R} \;, 
$$

где $\mathbf{R}$ верхняя треугольная матрица. Так как каждая из матриц $\mathbf{H}_k$ ортогональна, таковым будет и их произведение. Обратная от ортогональной также есть матрица ортогональная. Таким образом, алгоритм создает $\mathrm{QR}$ - разложение матрицы $\mathbf{A}$.

Напишите функцию, получающую прямоугольную матрицу $A$ и возвращающую матрицы $Q$ и $R$ --- компоненты $QR$-разложения $A$.


<!-- grade_id: cell-d5a63eec70d4d566, points: 0, grade: false, solution: true -->
```python
def qr_decomp(a):
    """Вычисляет QR - разложение матрицы.
    
    Parameters
    ----------
    a : ndarray, shape(m, n)
        Введенная матрица
    
    Returns
    -------
    q : ndarray, shape(m, m)
        Ортогональная матрица
    r : ndarray, shape(m, n)
        Верхняя треугольная матрица
        
    Examples
    --------
    >>> a = np.random.random(size=(3, 5))
    >>> q, r = qr_decomp(a)
    >>> np.testing.assert_allclose(q @ r, a, atol=1e-14)
    
    """
    A = np.array(a, copy=True, dtype=float)
    m, n = A.shape
    iter = min(m,n)
    #print(iter)

    QT = np.eye(m)
    Q = np.eye(m)

    for i in range(0,iter):
      #print("Aaf",A)
      v1, q = householder(A[:,0])
      A=q @ A
      qa=np.eye(m)
      qa[i:,i:]=q
      QT=qa @ QT
      Q=Q @ (np.transpose(qa))

      #print("Abef",A)
      A = A[1:,1:]
      #print("Aaf",A)

    R=QT@a
    return Q, R
```

```python
a = np.array([[1,2,3],[4,5,6],[7,8,9],[11,12,13]])

q, r = qr_decomp(a)
print(q@r)
```

Output:
```
[[ 1.  2.  3.]
 [ 4.  5.  6.]
 [ 7.  8.  9.]
 [11. 12. 13.]]
```

```python
    a = np.array([[1,2,3],[4,5,6],[7,8,9],[11,12,13]])


    A = np.array(a, copy=True, dtype=float)
    m, n = A.shape
    iter = min(m,n)
    #print(iter)

    QT = np.eye(m)
    Q = np.eye(m)

    for i in range(0,iter):
      #print("Aaf",A)
      v1, q = householder(A[:,0])
      A=q @ A
      qa=np.eye(m)
      qa[i:,i:]=q
      QT=qa @ QT
      Q=Q @ (np.transpose(qa))

      print("Abef",A)
      A = A[1:,1:]
      print("Aaf",A)

    R=QT@a
    print(Q@R)
```

Output:
```
Abef [[1.36747943e+01 1.53567209e+01 1.70386475e+01]
 [9.99200722e-16 7.84792855e-01 1.56958571e+00]
 [1.44328993e-15 6.23387496e-01 1.24677499e+00]
 [1.99840144e-15 4.08180351e-01 8.16360702e-01]]
Aaf [[0.78479285 1.56958571]
 [0.6233875  1.24677499]
 [0.40818035 0.8163607 ]]
Abef [[ 1.08218436e+00  2.16436873e+00]
 [-4.15678427e-16  8.77736390e-16]
 [-2.19178961e-16 -8.94737239e-16]]
Aaf [[ 8.77736390e-16]
 [-8.94737239e-16]]
Abef [[ 1.25338577e-15]
 [-2.26758202e-31]]
Aaf []
[[ 1.  2.  3.]
 [ 4.  5.  6.]
 [ 7.  8.  9.]
 [11. 12. 13.]]
```

```python
A = np.array([[1,2,3],[4,5,6],[7,8,9]])

Q = np.eye(3)
Qad=np.eye(3)


#print(householder(A[:,0]))
v1, q = householder(A[:,0])
Q = q @ Q
Qad=Qad @ (np.transpose(q))
print("q\n",q)
print("Q\n",Q)

A1 = q @ A
print("A1bef\n",A1)

A1 = A1[1:,1:]
print("A1\n",A1)
v1, q = householder(A1[:,0])
A2 = q @ A1
print("q\n",q)
qa=np.eye(3)
qa[1:,1:]=q
#print(qa)
Q = qa @ Q
Qad=Qad @ (np.transpose(qa))

print("A2\n",A2)

A2 = A2[1:,1:]
print(A2[:,0])
v1, q = householder(A2[:,0])
print("q\n",q)
A3 = q @ A2

print("A3\n",A3)

qa=np.eye(3)
qa[2:,2:]=q
print(qa)
Q = qa @ Q
Qad=Qad @ (np.transpose(qa))
print(Q)
R=Q@A

print(Qad@R)
```

Output:
```
q
 [[ 0.12309149  0.49236596  0.86164044]
 [ 0.49236596  0.72354671 -0.48379326]
 [ 0.86164044 -0.48379326  0.1533618 ]]
Q
 [[ 0.12309149  0.49236596  0.86164044]
 [ 0.49236596  0.72354671 -0.48379326]
 [ 0.86164044 -0.48379326  0.1533618 ]]
A1bef
 [[8.12403840e+00 9.60113630e+00 1.10782342e+01]
 [4.44089210e-16 7.32119416e-01 1.46423883e+00]
 [8.88178420e-16 5.31208978e-01 1.06241796e+00]]
A1
 [[0.73211942 1.46423883]
 [0.53120898 1.06241796]]
q
 [[ 0.80938847  0.58727362]
 [ 0.58727362 -0.80938847]]
A2
 [[9.04534034e-01 1.80906807e+00]
 [4.90360848e-17 8.16953540e-16]]
[8.1695354e-16]
q
 [[1]]
A3
 [[8.1695354e-16]]
[[1. 0. 0.]
 [0. 1. 0.]
 [0. 0. 1.]]
[[ 0.12309149  0.49236596  0.86164044]
 [ 0.90453403  0.30151134 -0.30151134]
 [-0.40824829  0.81649658 -0.40824829]]
[[1. 2. 3.]
 [4. 5. 6.]
 [7. 8. 9.]]
```

```python
# можете запустить данную операцию для бооее сжатого вывода: нули вместо `1e-16` и т.д.
np.set_printoptions(suppress=True)
```

<!-- grade_id: cell-c3391a61cbdadb2f, points: 3, grade: true, solution: false -->
```python
rndm = np.random.RandomState(1234)
a = rndm.uniform(size=(5, 3))
aa = a.copy()

q, r = qr_decomp(a)

# check that `qr_decomp` leaves `a` intact
assert_allclose(a, aa, atol=1e-16)

# тестируем, что Q ортогональна
assert_allclose(q @ q.T, np.eye(5), atol=1e-14)

# проверяем разложение
assert_allclose(q @ r, a, atol=1e-14)
```

Теперь сравните ваше разложение с разложением, полученным библиотечной функцией (содержащей соответствующие функции библиотеки LAPACK)

```python
from scipy.linalg import qr
qq, rr = qr(a)

assert_allclose(np.dot(qq, rr), a,
                atol=1e-14)
```

Проверьте, согласуются ли ваши матрицы `q` и `r` с `qq` и `rr`. Объясните результат.
Оставьте пояснения в этой ячейке.


### Часть III. Безматричная реализация.

Отметим необычную структуру матрицы Хаусхолдера: матрица поворота $\mathbf{H}$ полностью характеризуется вектором отражения $\mathbf{u}$. Заметим, также, что вычислительная сложность операции отражения матрицы сильно зависит от порядка операций:
$$ \left( \mathbf{u} \mathbf{u}^T \right) \mathbf{A} \qquad \textrm{is } O(m^2 n)\;, $$

тогда как $$ \mathbf{u} \left( \mathbf{u}^T \mathbf{A} \right) \qquad \textrm{is } O(mn) $$

Таким образом, следует избегать формирований матриц $\mathbf{H}$. Вместо этого можно сохранять сами векторы отражения $\mathbf{u}$ и производить умножение произвольной матрицы на $\mathbf{Q}^T$, например, как отдельную функцию (класс).

Напишите функцию, выполняющую QR - разложение матрицы без формирования матриц $\mathbf{H}$ и возвращающую матрицу $\mathbf{R}$, а также вектора отражений Хаусхолдера.


Напишите вторую функцию, которая использует вектора отражений, полученных предыдущей функцией, для вычисления произведения $Q^T B$ для заданной матрицы B. Убедитесь, что вы добавили достаточно комментариев, следующих вашим выкладкам. 



<!-- grade_id: cell-5360f8930eabe0b4, points: 0, grade: false, solution: true -->
```python
def qr_nomatrix(a):
    """Form QR decomposition of `a` via successive Householder reflections.
    
    Parameters
    ----------
    a : ndarray
        Input matrix
    
    Returns
    -------
    R : ndarray
        Upper triangular matrix of the QR decomposition
    
    U : ndarray
        Columns store successive Householder reflectors: `U[j:, j]` stores
        the Householder reflector for reducing the `j`-th column.
        
    See Also
    --------
    householder_apply : apply Householder reflectors stored in `U` to `a`.
    
    """
    A = np.array(a, copy=True, dtype=float)
    m, n = A.shape
    iter = min(m,n)
    print(m,n)

    QT = np.eye(m)
    Q = np.eye(m)
    U = np.eye(m)
    print(U)

    for i in range(0,iter):

      vec = np.asarray(A[:,0], dtype=float)
      if vec.ndim != 1:
          raise ValueError("vec.ndim = %s, expected 1" % vec.ndim)

      if np.size(vec)==1 and np.linalg.norm(vec)<1e-8:
        q = [[1]]
        u = np.asarray([0])
      else:
        x = vec
        y = np.zeros(np.size(vec))
        y[0] = np.linalg.norm(vec)
        u = (x - y)/np.linalg.norm(x - y)
        print("u",u)
      U[i:, i] = u
      print("U",U)


      q=np.eye(np.size(vec))-2*(u[np.newaxis, :].T @ u[np.newaxis, :] )


      A=q @ A
      qa=np.eye(m)
      qa[i:,i:]=q
      QT=qa @ QT
      Q=Q @ (np.transpose(qa))

      A = A[1:,1:]

    R=QT@a
    return R, U

    
def householder_apply(b, uu):
    """Apply the Householder reflectors stored in `uu` to `b`.
    
    The result is equivalent to
    >>> r, q = qr_decomp(a)
    >>> q.T @ b
    
    Parameters
    ----------
    b : ndarray
        Input matrix
    uu : ndarray
        Householder reflectors: `U[j:, j]` is the reflection vector
        for transforming the `j`-th column of `a`.
        
    Returns
    -------
    r : ndarray
        The result of applying the reflectors to `b`. Equivalent to
        computing `q.T @ b`.

    See Also
    --------
    qr_decomp
    
    """
    A = np.array(b, copy=True, dtype=float)
    m, n = A.shape
    iter = min(m,n)

    QT = np.eye(m)
    Q = np.eye(m)

    for i in range(0,iter):
      u = U[i:, i]

      vec = np.asarray(A[:,0], dtype=float)

      q=np.eye(np.size(vec))-2*(u[np.newaxis, :].T @ u[np.newaxis, :] )
      A=q @ A
      qa=np.eye(m)
      qa[i:,i:]=q
      QT=qa @ QT
      Q=Q @ (np.transpose(qa))

      A = A[1:,1:]

    R=QT@b
    return R
    
```

<!-- grade_id: cell-68363e2462de3a8b, points: 3, grade: true, solution: false -->
```python
rndm = np.random.RandomState(1223)

a = rndm.uniform(size=(5, 5))
R1, U = qr_nomatrix(a)
R2 = householder_apply(a, U)
R_lib = qr(a)[1]
assert_allclose(R1, R2, atol=1e-14)

```

Output:
```
5 5
[[1. 0. 0. 0. 0.]
 [0. 1. 0. 0. 0.]
 [0. 0. 1. 0. 0.]
 [0. 0. 0. 1. 0.]
 [0. 0. 0. 0. 1.]]
u [-0.51706743  0.1189934   0.38529641  0.72128983  0.22309079]
U [[-0.51706743  0.          0.          0.          0.        ]
 [ 0.1189934   1.          0.          0.          0.        ]
 [ 0.38529641  0.          1.          0.          0.        ]
 [ 0.72128983  0.          0.          1.          0.        ]
 [ 0.22309079  0.          0.          0.          1.        ]]
u [-0.71706866 -0.16784218 -0.19207851  0.64865043]
U [[-0.51706743  0.          0.          0.          0.        ]
 [ 0.1189934  -0.71706866  0.          0.          0.        ]
 [ 0.38529641 -0.16784218  1.          0.          0.        ]
 [ 0.72128983 -0.19207851  0.          1.          0.        ]
 [ 0.22309079  0.64865043  0.          0.          1.        ]]
u [-0.66214187 -0.2285334   0.71368104]
U [[-0.51706743  0.          0.          0.          0.        ]
 [ 0.1189934  -0.71706866  0.          0.       
```

<!-- grade_id: cell-159c6803d580116d, points: 3, grade: true, solution: false -->
```python
# Check consistency with the scipy library decomposition. Allow for sign differences

conds = [np.allclose(R2[i, :], R_lib[i, :], atol=1e-14) or
         np.allclose(R2[i, :], -R_lib[i, :], atol=1e-14) for i in range(5)]
assert False not in conds

```

<!-- grade_id: cell-cfb9d263bcd923ed, points: 3, grade: true, solution: false -->
```python
# More testing here, keep this cell intact

```

```python

```