# Интерполяция при помощи полиномов Лагранжа.

<!-- grade_id: cell-d459dee4d09527c4, points: 0, grade: false, solution: true -->
```python
import numpy as np

class LagrangeInterpolator:
    """Lagrange interpolating polynomial.
    
    Given a set of pairs ``(x_k, y_k)``, construct 
    a Lagrange polynomial ``f(x)``, such that
    
    .. math::

        f(x_k) = y_k   for k =0, ..., n-1
    
    Parameters
    ----------
    xk : array_like, shape(n,)
        Abscissas
    yk : array_like, shape(n,)
        Ordinates
    
    Attributes
    ----------
    __call__
    
    """
    def __init__(self, xk, yk):
        self.xk = np.asarray(xk, dtype=float)
        self.yk = np.asarray(yk, dtype=float)
        
    def __call__(self, x):
        """Evaluate the interpolator at a given point.
        
        Parameters
        ----------
        x : float
        
        Returns
        -------
        the value of the interpolator at ``x``.
        """
        ### BEGIN SOLUTION
        def lagrange_k(xval, xk, k):
            n = len(xk)
            num, den = 1.0, 1.0
            for j in range(n):
                if j == k:
                    continue          
                den *= (xk[k] - xk[j])
                num *= (xval - xk[j])
            return num / den

        return sum(self.yk[k] * lagrange_k(x, self.xk, k) for k in range(len(self.xk)))
        ### END SOLUTION
```

<!-- grade_id: cell-02e34e05d317dfb2, points: 5, grade: true, solution: false -->
```python
def runge_func(x, a=25):
    return 1.0 / (1.0 + a*x**2)

xx = np.linspace(-2, 2, 21)
yy = runge_func(xx)

lagr = LagrangeInterpolator(xx, yy)

from numpy.testing import assert_allclose

assert_allclose(yy,
                [lagr(xval) for xval in xx],
                atol=1e-14)
```

Рассмотрим функцию Рунге $1/(1 + 25x^2)$. Интерполируйте данную функцию на интервале $x\in [-2, 2]$, используя полиномы Лагранжа степени $m$, где $m=3, 5, 7, 11$. Используйте равномерную сетку. Постройте результат интерполяции вместе с $f(x)$.

Повторите предыдущее задание используя узлы Чебышева. Сравните качество интерполяции на равномерной сетке и на узлах Чебышева.


```python
# YOUR CODE AND COMMENTS HERE (5 pts, manually graded)
```