Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Vladislav Aleksandrov"
COLLABORATORS = "-"
```

---

```python
import numpy as np

from numpy.testing import assert_allclose
```


### Часть I. Постройте отражение Хаусхолдера для вектора.

Дан вектор $\mathbf{x}$ и плоскость, заданная вектором нормали $\mathbf{u}$. Преобразование Хаусхолдера отражает $\mathbf{x}$ относительно плоскости.

Матрица преобразований Хаусхолдера:
$$ \mathbf{H} = \mathbf{1} - 2 \mathbf{u} \mathbf{u}^T $$

Если даны два вектора $\mathbf{x}$ и $\mathbf{y}$ одинаковой длины, поворот, преобразующий $\mathbf{x}$ в $\mathbf{y}$ называется преобразованием Хаусхолдера с
$$ \mathbf{u} = \frac{\mathbf{x} - \mathbf{y}}{\left|\mathbf{x} - \mathbf{y}\right|} $$

Напишите функцию, преобразующую заданный вектор $\mathbf{x} = (x_1, \dots, x_n)$ в $\mathbf{y} = (\left|\mathbf{x}\right|, 0, \dots, 0)^T$, используя преобразование Хаусхолдера.


<!-- grade_id: cell-d8d82542e794c193, points: 0, grade: false, solution: true -->
```python
def householder(vec):
    """Создает отражение Хаусхолдера, преобразующее 2-ю, 3-ю и т.д. компоненты вектора vec в нули.
    
    Parameters
    ----------
    vec : array-like of floats, shape (n,)
        Введенный вектор
    
    Returns
    -------
    outvec : array of floats, shape (n,)
        Преобразованный вектор, причем ``outvec[1:]==0`` и ``|outvec| == |vec|``
    H : array of floats, shape (n, n)
        Ортогональная матрица отражений Хаусхолдера
    """
    vec = np.asarray(vec, dtype=float)
    if vec.ndim != 1:
        raise ValueError("vec.ndim = %s, expected 1" % vec.ndim)

    if vec[0] > 0:
      sum = ((vec[1:])**2).sum()
      norm = -sum / (vec[0] + np.linalg.norm(vec))
    else:
      norm = vec[0] - np.linalg.norm(vec)
    
    new_vec = vec.copy()
    new_vec[0] = norm
    if (np.linalg.norm(new_vec) == 0):
      return new_vec, np.eye(len(new_vec))
    u = np.transpose([(new_vec) / np.linalg.norm(new_vec)])
    H = np.eye(len(u)) - 2*(u @ u.T)
    return H @ vec, H
```

Протестируйте свою функцию на следующих примерах:

<!-- grade_id: cell-7364c2b9de6bb44c, points: 3, grade: true, solution: false -->
```python
v = np.array([1, 2, 3])
v1, h = householder(v)

assert_allclose(h @ v1, v, atol=1e-14)
assert_allclose(h @ v, v1, atol=1e-14)

assert_allclose(v1[1:], 0, atol=1e-14)

assert_allclose(h @ h.T, np.eye(3), atol=1e-14)

###BEGIN HIDDEN TESTS
rndm = np.random.RandomState(1223)
vec = rndm.uniform(size=21)
v1, h = householder(vec)

assert_allclose(h @ v1, vec, atol=1e-14)
assert_allclose(h @ vec, v1, atol=1e-14)
assert_allclose(v1[1:], 0, atol=1e-14)

assert_allclose(h @ h.T, np.eye(h.shape[0]), atol=1e-14)
###END HIDDEN TESTS
```

### Part II. Вычисление $\mathrm{QR}$ - разложения матрицы.

Дана прямоугольная $m\times n$ матрица $\mathbf{A}$. Выполните отражение Хаусхолдера матрицы $\mathbf{H}_1$, преобразующее первый столбец матрицы $\mathbf{A}$ (назовем результат $\mathbf{A}^{(1)}$)

$$ 
\mathbf{H}_1 \mathbf{A} %\begin{pmatrix} \times &amp; \times &amp; \times &amp; \dots &amp; \times \\ 0 &amp; \times &amp; \times &amp; \dots &amp; \times \\ 0 &amp; \times &amp; \times &amp; \dots &amp; \times \\ &amp;&amp; \dots&amp;&amp; \\ 0 &amp; \times &amp; \times &amp; \dots &amp; \times \\ \end{pmatrix}
%
\equiv \mathbf{A}^{(1)}\;. 
$$

Теперь рассмотрим нижнюю правую часть матрицы $\mathbf{A}^{(1)}$ и выполним преобразование Хаусхолдера, действующее на 2 столбец $\mathbf{A}$:
$$
\mathbf{H}_2 \mathbf{A}^{(1)} % \begin{pmatrix} \times &amp; \times &amp; \times &amp; \dots &amp; \times \\ 0 &amp; \times &amp; \times &amp; \dots &amp; \times \\ 0 &amp; 0 &amp; \times &amp; \dots &amp; \times \\ &amp;&amp; \dots&amp;&amp; \\ 0 &amp; 0 &amp; \times &amp; \dots &amp; \times \\ \end{pmatrix}
%
\equiv \mathbf{A}^{(2)} \;. $$

Повторяя процесс $n$ раз, получим
$$
\mathbf{H}_{n} \cdots \mathbf{H}_2 \mathbf{H}_1 \mathbf{A} = \mathbf{R} \;, 
$$

где $\mathbf{R}$ верхняя треугольная матрица. Так как каждая из матриц $\mathbf{H}_k$ ортогональна, таковым будет и их произведение. Обратная от ортогональной также есть матрица ортогональная. Таким образом, алгоритм создает $\mathrm{QR}$ - разложение матрицы $\mathbf{A}$.

Напишите функцию, получающую прямоугольную матрицу $A$ и возвращающую матрицы $Q$ и $R$ --- компоненты $QR$-разложения $A$.


<!-- grade_id: cell-d5a63eec70d4d566, points: 0, grade: false, solution: true -->
```python
def qr_decomp(a):
    """Вычисляет QR - разложение матрицы.
    
    Parameters
    ----------
    a : ndarray, shape(m, n)
        Введенная матрица
    
    Returns
    -------
    q : ndarray, shape(m, m)
        Ортогональная матрица
    r : ndarray, shape(m, n)
        Верхняя треугольная матрица
        
    Examples
    --------
    >>> a = np.random.random(size=(3, 5))
    >>> q, r = qr_decomp(a)
    >>> np.testing.assert_allclose(q @ r, a, atol=1e-14)
    
    """
    a1 = np.array(a, copy=True, dtype=float)
    m, n = a1.shape
    v, H = householder(a1[:,0])
    for i in range(1,n):
      v, h = householder(a1[i:,i])
      new_h = np.eye(H.shape[0])
      new_h[i:,i:] = h
      H = H @ new_h
    return H.T, H @ a

```

```python
# можете запустить данную операцию для бооее сжатого вывода: нули вместо `1e-16` и т.д.
np.set_printoptions(suppress=True)
```

<!-- grade_id: cell-c3391a61cbdadb2f, points: 3, grade: true, solution: false -->
```python
rndm = np.random.RandomState(1234)
a = rndm.uniform(size=(5, 3))
aa = a.copy()

q, r = qr_decomp(a)

# check that `qr_decomp` leaves `a` intact
assert_allclose(a, aa, atol=1e-16)

# тестируем, что Q ортогональна
assert_allclose(q @ q.T, np.eye(5), atol=1e-14)

# проверяем разложение
assert_allclose(q @ r, a, atol=1e-14)
```

Теперь сравните ваше разложение с разложением, полученным библиотечной функцией (содержащей соответствующие функции библиотеки LAPACK)

```python
from scipy.linalg import qr
qq, rr = qr(a)

assert_allclose(np.dot(qq, rr), a,
                atol=1e-14)
```

Проверьте, согласуются ли ваши матрицы `q` и `r` с `qq` и `rr`. Объясните результат. Объяснение:<br>
Результат различается т.к. QR разложение не единствено.


### Часть III. Безматричная реализация.

Отметим необычную структуру матрицы Хаусхолдера: матрица поворота $\mathbf{H}$ полностью характеризуется вектором отражения $\mathbf{u}$. Заметим, также, что вычислительная сложность операции отражения матрицы сильно зависит от порядка операций:
$$ \left( \mathbf{u} \mathbf{u}^T \right) \mathbf{A} \qquad \textrm{is } O(m^2 n)\;, $$

тогда как $$ \mathbf{u} \left( \mathbf{u}^T \mathbf{A} \right) \qquad \textrm{is } O(mn) $$

Таким образом, следует избегать формирований матриц $\mathbf{H}$. Вместо этого можно сохранять сами векторы отражения $\mathbf{u}$ и производить умножение произвольной матрицы на $\mathbf{Q}^T$, например, как отдельную функцию (класс).

Напишите функцию, выполняющую QR - разложение матрицы без формирования матриц $\mathbf{H}$ и возвращающую матрицу $\mathbf{R}$, а также вектора отражений Хаусхолдера.


Напишите вторую функцию, которая использует вектора отражений, полученных предыдущей функцией, для вычисления произведения $Q^T B$ для заданной матрицы B. Убедитесь, что вы добавили достаточно комментариев, следующих вашим выкладкам. 



<!-- grade_id: cell-5360f8930eabe0b4, points: 0, grade: false, solution: true -->
```python
def qr_nomatrix(a):
    """Form QR decomposition of `a` via successive Householder reflections.
    
    Parameters
    ----------
    a : ndarray
        Input matrix
    
    Returns
    -------
    R : ndarray
        Upper triangular matrix of the QR decomposition
    
    U : ndarray
        Columns store successive Householder reflectors: `U[j:, j]` stores
        the Householder reflector for reducing the `j`-th column.
        
    See Also
    --------
    householder_apply : apply Householder reflectors stored in `U` to `a`.
    
    """
    R = np.copy(a)
    m, n = R.shape
    U = np.empty((m, n))
    for i in range(n):
        vec = np.array(R[i:, i])

        if vec[0] > 0:
          sum = ((vec[1:])**2).sum()
          norm = -sum / (vec[0] + np.linalg.norm(vec))
        else:
          norm = vec[0] - np.linalg.norm(vec)

        new_vec = vec.copy()
        new_vec[0] = norm
        if (np.linalg.norm(new_vec) == 0):
          u = np.zeros(len(nev_vec))
        else:
          u = (new_vec) / np.linalg.norm(new_vec)

        U[i:, i] = u
        R[i:, i:] -= 2 * np.outer(u, (u @ R[i:, i:])) 
    return R, U



    
def householder_apply(b, uu):
    """Apply the Householder reflectors stored in `uu` to `b`.
    
    The result is equivalent to
    >>> r, q = qr_decomp(a)
    >>> q.T @ b
    
    Parameters
    ----------
    b : ndarray
        Input matrix
    uu : ndarray
        Householder reflectors: `U[j:, j]` is the reflection vector
        for transforming the `j`-th column of `a`.
        
    Returns
    -------
    r : ndarray
        The result of applying the reflectors to `b`. Equivalent to
        computing `q.T @ b`.

    See Also
    --------
    qr_decomp
    
    """
    A = np.array(b)
    n = uu.shape[1]

    for i in range(n):
      u = np.transpose([uu[i:, i]])
      A[i:, i:] -= 2 * u @ (u.T @ A[i:, i:])
    return A
    
```

<!-- grade_id: cell-68363e2462de3a8b, points: 3, grade: true, solution: false -->
```python
rndm = np.random.RandomState(1223)

a = rndm.uniform(size=(5, 5))
R1, U = qr_nomatrix(a)
R2 = householder_apply(a, U)
R_lib = qr(a)[1]
assert_allclose(R1, R2, atol=1e-14)

###BEGIN HIDDEN TESTS

name2 = "Test 2: vertical rectangular matrix"
a2 = rndm.uniform(size=(5, 3))
R1, U = qr_nomatrix(a2)
R22 = householder_apply(a2, U)
R_lib2 = qr(a2)[1]
assert_allclose(R1, R22, atol=1e-14)


name3 = "Test 3: horizontal rectangular matrix"
a3 = rndm.uniform(size=(3, 5))
R1, U = qr_nomatrix(a3)
R23 = householder_apply(a3, U)
R_lib3 = qr(a3)[1]
assert_allclose(R1, R23, atol=1e-14)

###END HIDDEN TESTS
```

Output:
```
<ipython-input-10-8f74a7315d63> in <module>
     19 name3 = "Test 3: horizontal rectangular matrix"
     20 a3 = rndm.uniform(size=(3, 5))
---> 21 R1, U = qr_nomatrix(a3)
     22 R23 = householder_apply(a3, U)
     23 R_lib3 = qr(a3)[1]

<ipython-input-9-27278fe1aab2> in qr_nomatrix(a)
     27         vec = np.array(R[i:, i])
     28 
---> 29         if vec[0] > 0:
     30           sum = ((vec[1:])**2).sum()
     31           norm = -sum / (vec[0] + np.linalg.norm(vec))

IndexError: index 0 is out of bounds for axis 0 with size 0
```

<!-- grade_id: cell-159c6803d580116d, points: 3, grade: true, solution: false -->
```python
# Check consistency with the scipy library decomposition. Allow for sign differences

conds = [np.allclose(R2[i, :], R_lib[i, :], atol=1e-14) or
         np.allclose(R2[i, :], -R_lib[i, :], atol=1e-14) for i in range(5)]
assert False not in conds

###BEGIN HIDDEN TESTS
conds = [np.allclose(R22[i, :], R_lib2[i, :], atol=1e-14) or
         np.allclose(R22[i, :], -R_lib2[i, :], atol=1e-14) for i in range(5)]
assert False not in conds

conds = [np.allclose(R23[i, :], R_lib3[i, :], atol=1e-14) or
         np.allclose(R23[i, :], -R_lib3[i, :], atol=1e-14) for i in range(3)]
assert False not in conds
###END HIDDEN TESTS
```

Output:
```
<ipython-input-11-8bf97c9a7aa3> in <module>
     11 
     12 conds = [np.allclose(R23[i, :], R_lib3[i, :], atol=1e-14) or
---> 13          np.allclose(R23[i, :], -R_lib3[i, :], atol=1e-14) for i in range(3)]
     14 assert False not in conds
     15 ###END HIDDEN TESTS

<ipython-input-11-8bf97c9a7aa3> in <listcomp>(.0)
     11 
     12 conds = [np.allclose(R23[i, :], R_lib3[i, :], atol=1e-14) or
---> 13          np.allclose(R23[i, :], -R_lib3[i, :], atol=1e-14) for i in range(3)]
     14 assert False not in conds
     15 ###END HIDDEN TESTS

NameError: name 'R23' is not defined
```

<!-- grade_id: cell-cfb9d263bcd923ed, points: 3, grade: true, solution: false -->
```python
# More testing here, keep this cell intact

### BEGIN HIDDEN TESTS

rndm = np.random.RandomState(1234)
a = rndm.uniform(size=(5, 3))

q, r = qr_decomp(a)
rr, uu = qr_nomatrix(a)

# check consistency
assert_allclose(np.triu(r),
                np.triu(rr), atol=1e-14)

# $Q^T @ A = R$
assert_allclose(householder_apply(a, uu),
                r,
                atol=1e-14)

# act on a different matrix
b = rndm.uniform(size=(5, 3))
assert_allclose(householder_apply(b, uu),
                q.T @ b,
                atol=1e-14)

# END HIDDEN TESTS
```

Output:
```
~/anaconda3/lib/python3.7/site-packages/numpy/testing/_private/utils.py in assert_allclose(actual, desired, rtol, atol, equal_nan, err_msg, verbose)
   1531     header = 'Not equal to tolerance rtol=%g, atol=%g' % (rtol, atol)
   1532     assert_array_compare(compare, actual, desired, err_msg=str(err_msg),
-> 1533                          verbose=verbose, header=header, equal_nan=equal_nan)
   1534 
   1535 

~/anaconda3/lib/python3.7/site-packages/numpy/testing/_private/utils.py in assert_array_compare(comparison, x, y, err_msg, verbose, header, precision, equal_nan, equal_inf)
    844                                 verbose=verbose, header=header,
    845                                 names=('x', 'y'), precision=precision)
--> 846             raise AssertionError(msg)
    847     except ValueError:
    848         import traceback

AssertionError: 
Not equal to tolerance rtol=1e-07, atol=1e-14

Mismatched elements: 6 / 15 (40%)
Max absolute difference: 1.22282523
Max relative diffe
```

```python

```