Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Plaksin Anton"
COLLABORATORS = ""
```

---

```python
import numpy as np
from numpy.testing import assert_allclose
```

# 1. Степенная итерация

Дана квадратная $n \times n$ матрица $A$. Алгоритм степенной итерации позволяет найти максимальное собственное значение $\lambda$ матрицы $A$. Для начала необходимо задать начальное приближение собственного вектора $x_0$. 


Модифицированный алгоритм степенной итерации:

1. $y_{k+1} = A x_k$

2. $x_{k+1} = \cfrac{y_{k+1}}{||y_{k+1}||}$

3. $\lambda^{(k)} = (x_k, y_{k+1})$

где $k$ - номер итерации.

Напишите функцию, реализующую поиск максимального собственного значения заданной матрицы с помощью алгоритма степенной итерации.

<!-- grade_id: cell-3cfc4f5a0e04dc11, points: 0, grade: false, solution: true -->
```python
def power_iter(a, vec, eps=1e-3, maxiter=1000):
    """Реализует алгоритм степенной итерации поиска максимального собственного значения матрицы `a` 
        с начальным приближением собственного вектора vec.
    
    Parameters
    ----------
    a : array-like of floats, shape (n, n)
        Введённая матрица
    vec : array-like of floats, shape (n,)
        Введённое начальное приближение собственного вектора
    eps : float
        Введённая абсолютная точность
    maxiter : integer
        Максимальное количество итераций
    
    Returns
    -------
    lm : float
        Приближение максимального собственного значения
    outvec : array of floats, shape (n,)
        Приближение собственного вектора
    iters : integer
        Количество итераций
    """
    a = np.asarray(a, dtype=float)
    if a.ndim != 2:
        raise ValueError("a.ndim = %s, expected 2" % a.ndim)
    
    vec = np.asarray(vec, dtype=float)
    if vec.ndim != 1:
        raise ValueError("vec.ndim = %s, expected 1" % vec.ndim)
        
    if (a.shape[0] != a.shape[1]):
        raise ValueError("a is not square")
        
    if (a.shape[0] != vec.shape[0]):
        raise ValueError("vec should be of length %s" % a.shape[0])
    
    xk = np.array(vec, dtype=np.float64)
    yk1 = a @ xk
    xk1 = yk1 / np.linalg.norm(yk1)
    lk1 = xk @ yk1
    xk = xk1
    lk = lk1
    for k in range(maxiter):
        yk1 = a @ xk
        xk1 = yk1 / np.linalg.norm(yk1)
        lk1 = xk @ yk1
        if abs(lk1 - lk) < eps and np.linalg.norm(xk - xk1) < eps:
            return lk1, xk, k+1
        xk = xk1
        lk = lk1
    return lk1, xk, maxiter
```

Протестируйте свою функцию на следующих примерах:

<!-- grade_id: cell-d67c4ed20e0cc45a, points: 4, grade: true, solution: false -->
```python
A = np.asarray([[4, 34],[7, 1]])

np.random.seed(123)
x = np.random.random(size=2)

l, e, k = power_iter(A, x)
assert_allclose(A @ e, l * e, atol=1e-3)

l, e, k = power_iter(A, x, 1e-5)
assert_allclose(A @ e, l * e, atol=1e-5)

###BEGIN HIDDEN TESTS

A = np.asarray([[1, 2], [0, 1]])
x = np.asarray([1,1])

l, e, k = power_iter(A, x)
assert_allclose(A @ e, l * e, atol=1e-3)

l, e, k = power_iter(A, x, 1e-5)
assert_allclose(A @ e, l * e, atol=1e-5)

###END HIDDEN TESTS
```

# 2. Обратная итерация

Этот метод может использоваться, например, для поиска собственных векторов матрицы, если уже имеется достаточно хорошее приближение её максимального собственного значения.

Для реализации алгоритма необходимо задать начальные приближения для собственного вектора $x_0$ и собственного значения $\mu$.

На каждой итерации метода решается система уравнений:
$$(A - \lambda^{(k)} I) y_{k+1} = x_k,$$
приэтом $\lambda^{(0)} = \mu$. После этого за новое приближение собственного вектора принимается нормированный вектор $y_{k+1}$:
$$x_{k+1} = \cfrac{y_{k+1}}{||y_{k+1}||}.$$

<!-- grade_id: cell-f3b15fcf5f5b9c62, points: 0, grade: false, solution: true -->
```python
def inverse_iter(a, vec, mu, eps=1e-3, maxiter=1000):
    """Реализует алгоритм степенной итерации поиска максимального собственного значения матрицы a 
        с начальным приближением собственного вектора vec и начальным приближением собственного значения mu
    
    Parameters
    ----------
    a : array-like of floats, shape (n, n)
        Введённая матрица
    vec : array-like of floats, shape (n,)
        Введённое начальное приближение собственного вектора
    mu : float
        Введённое начальное приближение собственного значения
    eps : float
        Введённая точность 
    maxiter : integer
        Максимальное количество итераций
    
    Returns
    -------
    lm : float
        Приближение максимального собственного значения
    outvec : array of floats, shape (n,)
        Приближение собственного вектора
    iters : integer
        Количество итераций
    """
    a = np.asarray(a, dtype=float)
    if a.ndim != 2:
        raise ValueError("a.ndim = %s, expected 2" % a.ndim)
    
    vec = np.asarray(vec, dtype=float)
    if vec.ndim != 1:
        raise ValueError("vec.ndim = %s, expected 1" % vec.ndim)
        
    if (a.shape[0] != a.shape[1]):
        raise ValueError("a is not square")
        
    if (a.shape[0] != vec.shape[0]):
        raise ValueError("vec should be of length %s" % a.shape[0])

    xk = vec
    lk = mu
    for k in range(maxiter + 1):
        yk1 = np.linalg.solve((a - lk * np.eye(a.shape[0])), xk)
        xk1 = yk1 / np.linalg.norm(yk1)
        lk1 = (a @ xk) @ xk
        if abs(lk1 - lk) < eps and np.linalg.norm(xk1 - xk) < eps:
            return lk1, xk1, k
        xk = xk1
        lk = lk1
    return lk, xk, maxiter
```

Протестируйте свою функцию на следующих примерах:

<!-- grade_id: cell-014b5a8fe53ecfdf, points: 3, grade: true, solution: false -->
```python
A = np.asarray([[4, 34],
                [7, 1]])
x = np.random.random(size=2)

l, e, k = inverse_iter(A, x, 1.)
assert_allclose(A @ e, l * e, atol=1e-3)

l, e, k = inverse_iter(A, x, 100.)
assert_allclose(A @ e, l * e, atol=1e-3)

###BEGIN HIDDEN TESTS

A = np.asarray([[1, 2], [0, 1]])
x = np.asarray([1, 1])

l, e, k = inverse_iter(A, x, 0.5)
assert_allclose(A @ e, l * e, atol=1e-3)

l, e, k = inverse_iter(A, x, 50.)
assert_allclose(A @ e, l * e, atol=1e-3)

###END HIDDEN TESTS
```

Output:
```
~/anaconda3/lib/python3.7/site-packages/numpy/linalg/linalg.py in solve(a, b)
    397     signature = 'DD->D' if isComplexType(t) else 'dd->d'
    398     extobj = get_linalg_error_extobj(_raise_linalgerror_singular)
--> 399     r = gufunc(a, b, signature=signature, extobj=extobj)
    400 
    401     return wrap(r.astype(result_t, copy=False))

~/anaconda3/lib/python3.7/site-packages/numpy/linalg/linalg.py in _raise_linalgerror_singular(err, flag)
     95 
     96 def _raise_linalgerror_singular(err, flag):
---> 97     raise LinAlgError("Singular matrix")
     98 
     99 def _raise_linalgerror_nonposdef(err, flag):

LinAlgError: Singular matrix
```

## Пример

Рассмотрите матрицу 
$$ A = 
\begin{bmatrix}
3 & 1 & 0 & 0\\ 
1 & 2 & 0 & 1\\ 
0 & 0 & 1 & 1\\ 
0 & 1 & 1 & 1
\end{bmatrix}.$$

<!-- grade_id: cell-3fb3b96b54cbb97d, points: 0, grade: false, solution: true -->
```python
def A_and_x(seed=1234):
    """Возвращает матрицу `А` и случайный начальный вектор."""
    from numpy.random import default_rng

    rng = default_rng(seed)
    a = np.array([[3,1,0,0],[1,2,0,1],[0,0,1,1],[0,1,1,1]])
    vvals = rng.standard_normal(4)
    return a, vvals
```

Найдите максимальное собственное число этой матрицы тремя разными методами.

### 1) Степенная итерация

<!-- grade_id: cell-a48d281064127067, points: 1, grade: true, solution: false -->
```python
A, x = A_and_x()

l, e, k = power_iter(A, x, eps=1e-5)
assert_allclose(A @ e, l * e, atol=1.6e-2)

### BEGIN HIDDEN TESTS
A, x = A_and_x()
l, e, k = power_iter(A, x, eps=1e-5)
assert_allclose(A @ e, l * e, atol=1.6e-3)
### END HIDDEN TESTS
```

### 2) Обратная итерация с $\mu = 3.5$

<!-- grade_id: cell-80ce0d792279ba41, points: 1, grade: true, solution: false -->
```python
A, x = A_and_x()

l, e, k = inverse_iter(A, x, 3.5, eps=1e-5)
assert_allclose(A @ e, l * e, atol=1e-3)

### BEGIN HIDDEN TESTS
A, x = A_and_x()
l, e, k = inverse_iter(A, x, 3.5, eps=1e-5)
assert_allclose(A @ e, l * e, atol=1e-5)
### END HIDDEN TESTS
```

### 3) Обратная итерация с $\mu = 3.7$

<!-- grade_id: cell-9fdc8e167e855b76, points: 1, grade: true, solution: false -->
```python
A, x = A_and_x()

l, e, k = inverse_iter(A, x, 3.7, eps=1e-5)
assert_allclose(A @ e, l * e, atol=1e-3)

### BEGIN HIDDEN TESTS
A, x = A_and_x()
l, e, k = inverse_iter(A, x, 3.7, 1e-5)
assert_allclose(A @ e, l * e, atol=1e-5)
### END HIDDEN TESTS
```

Сколько шагов $k$ требуется в каждом случае для того, чтобы получить настоящий собственный вектор $x^*$ с точностью $||x^* - x_k||_2 < 10^{-3}$?

<!-- grade_id: cell-6e9a3b55d1fa9b7d, points: 0, grade: true, solution: true -->
```python
print(k, kk, kkk)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-11-75a6583158c5> in <module>
----> 1 print(k, kk, kkk)

NameError: name 'kk' is not defined
```